Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you add your code only in the place where it says `YOUR CODE HERE` or "YOUR ANSWER HERE". You must remove the line containing the `raise` command. Add your name and roll number below. If you have discussed any assignment problem with another student, mention the name/roll number of the student along with the problem number as a comment in the box below (using `(* comment *)`).

```OCaml
let name = ""
let rollno = ""
```

## Important notes about grading:

1. All code you submit must compile. Programs that do not compile will receive an automatic zero. If you are having trouble getting your assignment to compile, please contact the TAs or the instructor. 
2. Make sure you do not add any new cells. You must write your code in the designated place only. We use an autograder to grade the submissions, and if your code is not in it's designated place, or you have added a new cell, the autograder may not work correctly. In particular, executing the last cell of the file always adds a new cell at the end, so you must delete this newly added cell.
3. All assignments submitted after the deadline will be considered late, and will consume your grace days. 
4. Your code will also be tested on private testcases, with equal weightage for each public and private testcase. Make sure you test your code on your own testcases as well.

# Lambda Calculus Interpreter

In this assignment, you will implement a lambda calculus interpreter which can reduce any lambda term to its normal form. The abstract syntax tree (AST) for lambda expressions is the one that we have seen in class:


```ocaml
type expr = 
  | Var of string
  | Lam of string * expr
  | App of expr * expr
```


You are provided a parser function `parse_string` that converts a string to this AST and a printer function `string_of_expr` that converts the AST to string. For example, 

```OCaml
#use "init.ml"
open Syntax
let parse_string = Lambda_parse.parse_string
let string_of_expr = string_of_expr

let _ = parse_string "(\\x.x) (\\y.y)"
let _ = string_of_expr (App (Var "x",Lam("y",App(Var "y", Var "x"))))
```

You will need some helper functions to operate over sets. Since we have not studied set data structure in OCaml, we will use lists instead and implement set functionality on top of lists. **You can use any function from the OCaml list standard library for this assignment.** The documentation for OCaml list module is available in the [OCaml manual](https://v2.ocaml.org/api/List.html). **There is no requirement to make any function tail-recursive in this assignment.**

## Problem 1

Implement a function

```ocaml
mem : 'a -> 'a list -> bool
```

`mem e l` returns `true` if the element `e` is present in the list. Otherwise, it returns false. 

<!-- grade_id: cell-206225a8fc096ef3, points: 0, grade: false, solution: true -->
```OCaml
let mem e l = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-625148cc24b57f40, points: 5, grade: true, solution: false -->
```OCaml
(* 5 points *)
assert (mem "b" ["a";"b";"c"] = true);
assert (mem "x" ["a";"b";"c"] = false)
```

## Problem 2

Implement a function

```ocaml
remove_duplicates : 'a list -> 'a list
```

such that for every element `v` in the input list `l`, there is exactly one instance of `v` in the output list `remove_duplicates l`. Further, ordering of any two elements `v1` and `v2` in the output list is based on the ordering of the first instance of `v1` and `v2` in the input list.  For example, `remove_duplicates [1;2;2;3;1;1;1;4;4;2;2]= [1;2;3;4]`. `remove` from Assignment-1 might be handy here.


<!-- grade_id: cell-4dd11e969c04a997, points: 0, grade: false, solution: true -->
```OCaml
let rec remove_duplicates l = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-73687fe02c21e253, points: 5, grade: true, solution: false -->
```OCaml
(* 5 points *)
assert (remove_duplicates [1;2;2;3;1;1;1;4;4;2;2] = [1; 2; 3; 4]);
assert (remove_duplicates [1;1;1;1;1] = [1])
```

## Problem 3

Implement a function

```ocaml
union : string list -> string list -> string list
```

`union l1 l2` performs set union of strings in `l1` and `l2`. That is, the output list should contain exactly one instance of every string in the input lists `l1` and `l2`. Further, the elements in the output list must be lexicographically sorted. You can use the functions `List.sort` and `remove_duplicates` to implement union. Here is an example of using `List.sort`. 

```OCaml
assert (List.sort String.compare ["x";"a";"b";"m"] = ["a";"b";"m";"x"])
```

`List.sort` is a higher-order function that takes an input a binary `compare` function. For this problem, you can use the function `String.compare`, since the `union` function only operates over string lists.

<!-- grade_id: cell-e498f15f8517e45a, points: 0, grade: false, solution: true -->
```OCaml
let union l1 l2 = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-42c81525ae442418, points: 5, grade: true, solution: false -->
```OCaml
(* 5 points *)
assert (union ["a"; "c"; "b"] ["d"; "b"; "x"; "a"] = ["a"; "b"; "c"; "d"; "x"]);
assert (union ["a";"b"] ["c"] = ["a";"b";"c"])
```

## Substitution

At the heart of reducing lambda expressions is substitution. Recall from the lecture that substitution requires us to generate fresh variable names that are different from every other name used in the current context. We will use the following helper function to generate fresh names.

```OCaml
let r = ref 0 

let r' = r
                                                                                    
let fresh s =                                                                       
  let v = !r in                                                                     
  r := !r + 1;                                                                      
  s ^ (string_of_int v)
```

It uses mutability features of OCaml which we will study in later lectures. You can use the `fresh` function as follows:

```OCaml
let a0 = fresh "a"
let a1 = fresh "a"
```

That is, `fresh s` generates a new string obtained by appending `s` with a number. Every call to `fresh s` is guaranteed to append a different number to `s` compared to all the previous calls.  

## Problem 4

Implement a function 

```ocaml
fresh2: string -> string list -> string
```

such that `fresh2 s l` returns a version of string `s` that is guaranteed to be different from every string in `l`. For example, `fresh2 "a" ["a0";"a1";"b1";"a2"]` cannot return `a0`, `a1` or `a2`, but can return any other version (e.g. `a3`). You should use the `fresh` function in the definition of `fresh2`. See the testcases for more clarity.

<!-- grade_id: cell-a2baabb3cd58c967, points: 0, grade: false, solution: true -->
```OCaml
let rec fresh2 s vars = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-9e688c81a2e6f4bf, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (List.exists (fun s -> s = (fresh2 "a" ["a0";"a1";"b0"])) ["a0";"a1";"b0"] = false);
assert (List.exists (fun s -> s = (fresh2 "b" ["a0";"a1";"b0"])) ["a0";"a1";"b0"] = false)

```

## Problem 5

Implement a function 

```ocaml
free_variables: expr -> string list
```

that returns the free variables in the given lambda term. You can use the function ` union` defined earlier. The function `remove` from Assignment-1 can be used for removing an element from a set.

<!-- grade_id: cell-8ad0bb22f3fca6ed, points: 0, grade: false, solution: true -->
```OCaml
let rec free_variables e = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-67d38d859ec6173a, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (free_variables (parse_string "\\x.x") = []);
assert (free_variables (parse_string "\\x.y") = ["y"]);
assert (free_variables (parse_string "\\x.x (\\y. y)") = [])

```

## Problem 6

Implement the function

```ocaml
substitute : expr -> string -> expr -> expr
```

where `substitute e x v` does `e[v/x]`. Please use the precise definition of substitute from the lectures. For renaming `"x"` in `Lam("x",e)`, you must use a fresh variable which can be obtained by calling `fresh "x"`. 

<!-- grade_id: cell-c659e5c57a377290, points: 0, grade: false, solution: true -->
```OCaml
let rec substitute expr a b =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-d15558191ca8be99, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (alpha_equiv 
          (substitute (parse_string "\\y.x") "x" (parse_string "\\z.z w")) 
          (parse_string "λy.λz.z w"));
assert (alpha_equiv 
          (substitute (parse_string "\\x.x") "x" (parse_string "y"))
          (parse_string "λx.x"));
assert (alpha_equiv 
          (substitute (parse_string "\\x.y") "y" (parse_string "x"))
          (parse_string "λx0.x"))
```

## Problem 7

Implement the function

```ocaml
alpha_eq : expr -> expr -> bool
```

where `alpha_eq e1 e2` returns `true` if the lambda expressions `e1` and `e2` are $\alpha$-equivalent, `false` otherwise. Please use the precise definition of $\alpha$-equivalence from the lectures. You will need to use the `fresh2` function defined earlier in Problem 4.

<!-- grade_id: cell-eb52b8099ce39ff0, points: 0, grade: false, solution: true -->
```OCaml
let rec alpha_eq expr1 expr2 = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
```

<!-- grade_id: cell-d23597deb430fb31, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (alpha_eq (parse_string "\\x.x") (parse_string "\\y.y "));
assert (alpha_eq (parse_string "\\x. \\y. y") (parse_string "\\p. \\q. q"));
assert (alpha_eq (parse_string "\\x. \\y. \\z. x y z") (parse_string "\\y. \\z. \\x. y z x"))



```

## Problem 8

In this question, we will demonstrate Church-Rosser Theorem and uniqueness of $\beta$-normal forms. First, implement the function


```ocaml
beta_reduce_one_step : expr -> expr list
```

which performs **one** step of full $\beta$-reduction on its input lambda expression, and outputs a list of lambda-expressions which can be obtained after one step. If no step can be performed (i.e the input expression is in $\beta$-normal form), then the output will be an empty list. The lambda-expressions in the output can be in any order.

Next, implement the function

```ocaml
beta_reduce : expr -> expr list 
```

which performs full $\beta$-reduction on its input lambda expression until it reaches $\beta$-normal form. For this, you need to repeatedly invoke the `beta_reduce_one_step` function. Note that you need to explore all possible $\beta$-reductions. The output is a list of $\beta$-normal forms that you can achieve through all possible reduction paths (see the testcases for more clarity).

However, Church-Rosser theorem guarantees that you will achieve the same $\beta$-normal form through all reduction paths. Hence, the output should contain the same $\beta$-normal form repeated multiple times, with the number of repetitions being equal to the number of possible $\beta$-reduction paths.

Now, use `beta_reduce` to implement the function

```ocaml
beta_normal_form: expr -> expr
```

which returns the unique $\beta$-normal form of its input lambda expression.


<!-- grade_id: cell-bed846480aaf1da8, points: 0, grade: false, solution: true -->
```OCaml
let rec beta_reduce_one_step e = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
    
let beta_reduce e =    
        (* YOUR CODE HERE *)
        raise (Failure "Not implemented")
let beta_normal_form e = 
        (* YOUR CODE HERE *)
        raise (Failure "Not implemented")

```

<!-- grade_id: cell-41fe3cf7ee2fd5ee, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
let _ = assert (List.length (beta_reduce (parse_string "(\\ x.x x) ((\\x. y) z)")) = 3)
let _ = List.map (fun e -> assert (alpha_equiv e (parse_string "y y")))
    (beta_reduce (parse_string "(\\ x.x x) ((\\x. y) z)"))
let _ = assert (alpha_equiv (beta_normal_form (parse_string "(\\ x.x x) ((\\x. y) z)")) (parse_string "y y") )

let _ = assert (List.length (beta_reduce (parse_string "(\\ t. \\ f. t) x y")) = 1)
let _ = assert (alpha_equiv (beta_normal_form (parse_string "(\\ t. \\ f. t) x y")) (parse_string "x") )

let _ = assert (List.length (beta_reduce (parse_string "(\\ t. \\ f. t) (\\ x.x) y")) = 1)
let _ = assert (alpha_equiv (beta_normal_form (parse_string "(\\ t. \\ f. t) (\\ x.x) y")) (parse_string "\\ x.x") )

```

<!-- grade_id: cell-82ddc96e10e198dd, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
let zero = parse_string "\\f.\\x. x" 
let one = parse_string "\\f.\\x. f x" 
let two = parse_string "\\f.\\x. f (f x)" 
let three = parse_string "\\f.\\x. f (f (f x))" 

let plus = parse_string "λm. λn. λs. λz. m s (n s z)" 
let mult = parse_string "λm. λn. λs. λz. m (n s) z" 

let _ = assert (alpha_equiv (beta_normal_form (App (App (plus, one), two))) three)
let _ = assert (alpha_equiv (beta_normal_form (App (App (mult, one), three))) three)
```

## Problem 9

Implement the function 

```ocaml
beta_eq : expr -> expr -> bool
```

where `beta_eq e1 e2` returns `true` if the lambda expressions `e1` and `e2` are $\beta$-equivalent, `false` otherwise. Please use the precise definition of $\beta$-equivalence from the lectures. The `alpha_eq` and `beta_normal_form` functions defined in the previous problems might be useful.

<!-- grade_id: cell-c1b1375c4e0fcc83, points: 0, grade: false, solution: true -->
```OCaml
let beta_eq e1 e2 = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
  
```

<!-- grade_id: cell-3bd3e0e6187af09d, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
assert (beta_eq (parse_string "(\\ x.x x) ((\\x. y) z)") (parse_string "y y") );
assert (beta_eq (parse_string "(\\ t. \\ f. t) x y") (parse_string "x"));
assert (beta_eq (parse_string "(\\ t. \\ f. t) (\\ x.x) y") (parse_string "\\ y.y"))
```

## Problem 10

In this problem, we will repeat the steps in Problems 9,10 but for $\beta \eta$-reduction and $\beta \eta$-equivalence of lambda-expressions. Implement the following functions:

```ocaml
beta_eta_reduce_one_step : expr -> expr list
```

```ocaml
beta_eta_reduce : expr -> expr list 
```

```ocaml
beta_eta_normal_form: expr -> expr
```

```ocaml
beta_eta_eq : expr -> expr -> bool
```

Use the precise definition of $\beta \eta$-reduction and $\beta \eta$-equivalence from the lectures. You can use your implementations from the earlier problems.

<!-- grade_id: cell-8572530d8f782acf, points: 0, grade: false, solution: true -->
```OCaml
let rec beta_eta_reduce_one_step e = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")

let beta_eta_reduce e =     
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
    
let beta_eta_normal_form e = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")

let beta_eta_eq e1 e2 = 
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
```

<!-- grade_id: cell-33c21ddb14bf7695, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
let _ = assert (List.length (beta_eta_reduce (parse_string "\\x. f x")) = 1)
let _ = assert (alpha_equiv (beta_eta_normal_form (parse_string "\\x. f x")) (parse_string "f") )
let _ = assert (beta_eta_eq (parse_string "\\x. f x") (parse_string "f"))

let _ = assert (List.length (beta_eta_reduce (parse_string "(\\ x.x x) ((\\x. y) z)")) = 3)
let _ = List.map (fun e -> assert (alpha_equiv e (parse_string "y y")))
    (beta_eta_reduce (parse_string "(\\ x.x x) ((\\x. y) z)"))
let _ = assert (alpha_equiv (beta_eta_normal_form (parse_string "(\\ x.x x) ((\\x. y) z)")) (parse_string "y y") )

let _ = assert (List.length (beta_eta_reduce (parse_string "(\\ t. \\ f. t) (\\ x. g x) y")) = 3)
let _ = assert (beta_eta_eq  (parse_string "(\\ t. \\ f. t) (\\ x. g x) y") (parse_string "g") )

```