Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you add your code only in the place where it says `YOUR CODE HERE` or "YOUR ANSWER HERE". You must remove the line containing the `raise` command. Add your name and roll number below. If you have discussed any assignment problem with another student, mention the name/roll number of the student along with the problem number as a comment in the box below (using `(* comment *)`).

```OCaml
let name = ""
let rollno = ""
```

## Important notes about grading:

1. All code you submit must compile. Programs that do not compile will receive an automatic zero. If you are having trouble getting your assignment to compile, please contact the TAs or the instructor. 
2. Make sure you do not add any new cells. You must write your code in the designated place only. We use an autograder to grade the submissions, and if your code is not in it's designated place, or you have added a new cell, the autograder may not work correctly. In particular, executing the last cell of the file always adds a new cell at the end, so you must delete this newly added cell.
3. All assignments submitted after the deadline will be considered late, and will consume your grace days. 
4. Your code will also be tested on private testcases, with equal weightage for each public and private testcase. Make sure you test your code on your own testcases as well.

# Mutability and Modules

In this assignment, you will design and implement a couple of mutable data structures and operations on them.

## Problem 1

Implement structures `IntShowable` and `FloatShowable` that satisfy the following signature. Use `string_of_int: int -> string` and `string_of_float: float -> string` for the corresponding `string_of_t` functions.

```OCaml
module type Showable = sig
  type t
  val string_of_t : t -> string
end
```

<!-- grade_id: cell-158d909bf026ab70, points: 0, grade: false, solution: true -->
```OCaml
(* Implement modules IntShowable and FloatShowable below. DO NOT CHANGE THE SHARING CONSTRAINTS *)


module IntShowable:Showable with type t = int =
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
end

module FloatShowable: Showable with type t = float =
struct
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
end

```

<!-- grade_id: cell-e2c519b404ae59d8, points: 5, grade: true, solution: false -->
```OCaml
(* 5 points *)
assert (IntShowable.string_of_t 10 = "10");
assert (FloatShowable.string_of_t 0.0 = "0.")
```

## Problem 2

Implement a functor 

```ocaml
module MakeNode : 
  functor (C : Showable) -> 
    DoublyLinkedListNode with type content = C.t
``` 

where `DoublyLinkedListNode` is the following module type. Use the definition of a doubly linked list node from the lectures (however it will no longer be parameterized by a type variable). You need to ensure that the content field is also mutable.

```OCaml
module type DoublyLinkedListNode = sig
  type t
  (** The type of doubly linked list node *)
  
  type content
  (** The type of content stored in the doubly linked list node *)

  val create : content -> t 
  (** create a new doubly linked list node with [content] as the content. 
      The next and previous nodes are [None]. *)
  
  val get_next : t -> t option
  (** [get_next t] returns [Some t'] if [t'] is the successor node of [t]. 
      If [t] has no successor, then return [None] *)
      
  val get_prev : t -> t option
  (** [get_prev t] returns [Some t'] if [t'] is the predecessor node of [t]. 
      If [t] has no predecessor, then return [None] *)
      
  val get_content : t -> content
  (** [content t] returns the content [c] of node [t] *)
      
  val set_next : t -> t option -> unit
  (** [set_next t t'] updates the next node of [t] to be [t'] *)
  
  val set_prev : t -> t option -> unit
  (** [set_prev t t'] updates the prev node of [t] to be [t'] *)
  
  val set_content : t -> content -> unit
  (** [set_content t c] updates the content of [t] to be [c] *)
  
  val string_of_content : t -> string
  (** [string_of_content t] returns the string form of content stored in the node t *)
end
```

<!-- grade_id: cell-ef73360426263f29, points: 0, grade: false, solution: true -->
```OCaml
(* Implement the functor MakeNode. DO NOT CHANGE THE SHARING CONSTRAINTS *)


module MakeNode (C : Showable) : DoublyLinkedListNode with type content = C.t = struct
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
end
```

```OCaml
module IntNode = MakeNode(IntShowable)
```

<!-- grade_id: cell-33796bd0b9bbeaaa, points: 4, grade: true, solution: false -->
```OCaml
(* 4 points *)
let open IntNode in 
let i1 = create 1 in
assert (get_content i1 = 1);
```

<!-- grade_id: cell-9ac4ed1728b424da, points: 6, grade: true, solution: false -->
```OCaml
(* 6 points *)
let open IntNode in 
let i1,i2,i3 = create 1, create 2, create 3 in
set_next i1 @@ Some i2;
set_next i2 @@ Some i3;
set_prev i2 @@ Some i1;
set_prev i3 @@ Some i2;
let _ = assert (match get_next i1 with None -> false | Some i2 -> get_content i2 = 2) in
()
```

## Problem 3

Implement a functor 

```ocaml 
module MakeList : 
  functor (N : DoublyLinkedListNode) -> 
    DoublyLinkedList with type node = N.t 
                      and type content = N.content
```

where `DoublyLinkedList` is the following module type. Use the definition of doubly linked list from the lectures.


```OCaml
module type DoublyLinkedList = sig
  type t 
  (** The type of doubly linked list *)
  
  type node
  (** The type of a doubly linked list node *)
  
  type content
  (** The type of content stored in the [node] of doubly linked list *)
    
  val create : unit -> t 
  (** Creates a new (empty) doubly linked list *)
  
  val assign : t -> node option -> unit
  (** [assign t n] makes the node as the head node of the list. 
      The original contents of the list are dropped. *)
  
  val t_of_list : content list -> t
  (** [t_of_list l] returns a new doubly linked list with the list of 
      elements from the list [l]. *)
  
  val is_empty : t -> bool
  (** [is_empty l] return true if [t] is empty *)
  
  val first : t -> node option
  (** [first l] return [Some n] if the list is non-empty. Otherwise, return [None] *)
  
  val last : t -> node option
  (** [last l] returns the last node in the list. If the list is empty, return None**)
  
  val insert_first : t -> content -> node
  (** [insert_first l c] inserts a new doubly linked list node [n] as the first 
      node in [l] with [c] as the content. Returns [n]. *)
      
  val insert_after : node -> content -> node
  (** [insert_after n c] inserts a new node [n'] with content [c] after the node [n].
      Returns [n']. *)
      
  val insert_before : t -> node -> content -> node
  (** [insert_before n c] inserts a new node [n'] with content [c] before the node [n] in 
      the list t. Returns [n']. *)
  
  val remove : t -> node -> unit
  (** [remove l n] removes the node [n] from the list [l] *)
  
  val iter : t -> (content -> unit) -> unit
  (** [iter l f] applies [f] to each element of the list in the list order. *)
  
  val iter_node : t -> (node -> unit) -> unit
  (** [iter_node l f] applies [f] to each node of the list in the list order. 
      [f] may delete the current node or insert a new node after the current node. 
      In either case, [iter_node] is applied to the rest of the original list. *)
  
  val find_node : t -> (node -> bool) -> node option
  (** [find_node l f] applies the boolean function f on each node of the list in list order,
      and returns the first node in the list that satisfies f. If no node in the list satisfies f, 
      then it returns None.
  **)
      
  val string_of_t : t -> string
  (** [string_of_t (t_of_list [1;2;3])] returns the string "[1->2->3]".
      [string_of_t (t_of_list [])] returns the string "[]" *)

  val get_at_index : t -> int -> node option
  (** get_at_index l i returns node at index i. Assume that indexing starts from 0. If there is no node
  at the given index, return None.*)
  
  val size : t -> int
  (** returns the size of the list, i.e. the number of nodes in the list*)
end
```

<!-- grade_id: cell-37a66fdc4f8f3b4b, points: 0, grade: false, solution: true -->
```OCaml
(* Implement the functor MakeList. DO NOT CHANGE THE SHARING CONSTRAINTS *)


module MakeList (N : DoublyLinkedListNode) : DoublyLinkedList with type content = N.content 
                                                               and type node = N.t = 
struct
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
end


```

<!-- grade_id: cell-13adb3683de17ec0, points: 20, grade: true, solution: false -->
```OCaml
(* 20 points *)
let module L = MakeList(IntNode) in

let open L in
assert (is_empty (t_of_list []));

let l = t_of_list [1;2;3] in
assert (not @@ is_empty l);

let n = match first l with 
| None -> failwith "impossible" 
| Some n -> n
in
assert (1 = IntNode.get_content n);

let n = match last l with 
| None -> failwith "impossible" 
| Some n -> n
in
assert (3 = IntNode.get_content n);


L.(assign l (first (t_of_list [4;5;6])));
match first l with
| None -> failwith "impossible"
| Some n -> assert (IntNode.get_content n = 4);

let l = t_of_list [4;5;6] in
let n = get_at_index l 1 in
match n with
| None -> failwith "impossible"
| Some n -> assert (IntNode.get_content n = 5); let _ = insert_after n 7 in
assert (string_of_t l = "[4->5->7->6]")


```

<!-- grade_id: cell-77b4b373e05637e5, points: 5, grade: true, solution: false -->
```OCaml
(* 5 points *)
let module L = MakeList(IntNode) in

let open L in
let l = t_of_list [1;2;3] in
let n1 = match first l with
  | Some n1 -> n1
  | None -> failwith "impossible"
in
let n2 = match IntNode.get_next n1 with
  | Some n2 -> n2
  | None -> failwith "impossible"
in
remove l n2;
let n3 = match IntNode.get_next n1 with
  | Some n3 -> n3
  | None -> failwith "impossible"
in
let _ = assert (IntNode.get_content n3 = 3) in
remove l n1;
match first l with
| None -> failwith "impossible"
| Some n3 -> assert (IntNode.get_content n3 = 3)

```

<!-- grade_id: cell-ee91aa7878aa72ef, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
let module L = MakeList(IntNode) in
let open L in
let l = t_of_list [1;2;3] in
let s = ref 0 in
iter l (fun c -> s := !s + c);
assert (!s = 6);
iter_node l (fun n -> if IntNode.get_content n mod 2 == 0 then remove l n);
let n1 = match first l with
  | None -> failwith "impossible"
  | Some n -> n
in
assert (IntNode.get_content n1 = 1)

```

## Problem 4

In the next couple of problems, the goal is to implement a database using the doubly linked list module from the previous problem. A database stores a collection of *records* (different from OCaml records), where each record consists of a *key* and a *value*, which themselves are a collection of fields (also called *schema* in database terminology). For example, in a student database, the key would be roll number, while the value would consist of fields such as name, DOB, stream, semester, etc. The key may also consist of multiple fields, for example, in a database storing details about faculty advisor of students, the key would be both the student roll number and the faculty name. The point is that for every key instance, there would be at most one record with that key in the database (in database terminology, this is also called the primary key constraint).

First, we create module types `Key` and `Value`:


```OCaml
module type Key = sig
  type t
  val compare : t -> t -> int
  val string_of_t : t -> string
end

module type Value = Showable
```

`compare k1 k2` returns a value strictly greater than `0` if `k1 > k2`, it returns `0` if `k1 = k2` and returns a value strictly less than `0` if `k1 < k2`.  Depending on the schema, we will create different modules of types `Key` and `Value`.

Now, implement the functor:

```ocaml
module MakeRecord : 
  functor (K : Key) (V: Value) -> 
    Record with type key_type = K.t
           with type value_type = V.t 
``` 
 
where the `Record` module type is defined below:

```OCaml
module type Record = sig
  type key_type
  type value_type
  type t
  (* type of the record *)
  
  val create_record: key_type -> value_type -> t
  (* create k v creates a new record containing k and v *)
  
  val get_key : t -> key_type
  (* get_key r returns the key stored in r *)
  
  val get_value : t -> value_type
  (* get_value r returns the value stored in r *)
  
  val string_of_t : t -> string
  (* string_of_t r returns a string of the form "k:v" where k and v are the key and value stored in r. 
     Use string_ot_t functions of the key and value modules *)
end
```

<!-- grade_id: cell-408bcffc4a499832, points: 0, grade: false, solution: true -->
```OCaml
(* DO NOT CHANGE THE SHARING CONSTRAINTS *)

module MakeRecord (K : Key) (V: Value) : Record with type key_type = K.t
                                                with type value_type = V.t = struct
    (* YOUR CODE HERE *)
    raise (Failure "Not implemented")
end
```

<!-- grade_id: cell-7bb407937b146751, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
module StudentsKey:Key with type t = string = struct
  type t = string
  let compare = String.compare
  let string_of_t k = k 
end

type student_value_type = {name: string; hostel: string; cgpa: float}

module StudentsValue:Value with type t = student_value_type = struct
  type t = student_value_type
  let string_of_t v = v.name ^ ", " ^ v.hostel ^ ", " ^ (string_of_float v.cgpa)
end

module StudentsRecord = MakeRecord (StudentsKey) (StudentsValue)

let sk1 = "AE19S078";;
let sv1 = {name = "Pearce"; hostel = "Jamuna"; cgpa = 7.57};;
let sr1 = StudentsRecord.create_record sk1 sv1;;
assert (StudentsRecord.string_of_t sr1 = "AE19S078:Pearce, Jamuna, 7.57")
let sk2 = "AE23D016";;
let sv2 = {name = "Bobbe"; hostel = "Mandakini"; cgpa = 6.93};;
let sr2 = StudentsRecord.create_record sk2 sv2;;
assert (StudentsValue.string_of_t (StudentsRecord.get_value sr2)= "Bobbe, Mandakini, 6.93")

module FacultyKey:Key with type t = string = struct
  type t = string
  let compare = String.compare
  let string_of_t k = k 
end

type faculty_value_type = {name:string; designation:string; salary:float}

module FacultyValue:Value with type t = faculty_value_type = struct
  type t = faculty_value_type
  let string_of_t v = v.name ^ ", " ^ v.designation ^ ", " ^ (string_of_float v.salary)
end

module FacultyRecord = MakeRecord (FacultyKey) (FacultyValue)

let ck1 = "CS15";;
let cv1 = {name="Elvin"; designation="Asst. Professor"; salary=10000.0};;
let ck2 = "CS10";;
let cv2 = {name="Gaston"; designation="Assoc. Professor"; salary=20000.0};;

let cr1 = FacultyRecord.create_record ck1 cv1;;
assert(FacultyRecord.string_of_t cr1 = "CS15:Elvin, Asst. Professor, 10000.")
let cr2 = FacultyRecord.create_record ck2 cv2;;
assert(FacultyRecord.get_key cr2 = "CS10")


```

## Problem 5

Next, implement the function:

```ocaml
module MakeDatabase : 
  functor (K : Key) (V: Value) -> 
    Database with type key_type = K.t
             with type value_type = V.t 
             with type record_type = MakeRecord (K)(V).t
``` 

where the module type `Database` is defined below. You must use the doubly linked list module from Problem 3 to store the records. For this, you need to call the MakeList functor with an appropriate module as argument. Each node in the doubly linked list will store a record. Further, you need to ensure that records are stored in increasing order of their keys (using the `compare` function of the `Key` module). 

As part of the `Database` module, you will need to implement a number of methods for accessing and updating records  which are typically provided by a database. Since you will be internally using the `DoublyLinkedList` module to store the records, you will have to use the methods of `DoublyLinkedList` to access and update the appropriate nodes (or their content) of the linked list. 

```OCaml
exception PrimaryKeyViolation

module type Database = sig
  type key_type
  type value_type
  type record_type
  type t
  (* the type of the database. Must be the type t from the DoublyLinkedList module *)
  
  val create_database : unit ->  t
  (* Creates an empty database *)
  
  val insert_record : t -> key_type -> value_type -> unit
  (** insert_record d k v inserts a new record with key k and value v in d 
      in such a way that the list d should remain sorted. If there already exists a record in d with key k, 
      then you should generate a PrimaryKeyViolation exception**)
  
  val select_record : t -> key_type -> record_type option
  (** select_record d k returns the (unique) record in the list d with key k. 
      If no such record exists, it should return the string None. **)
      
  val select_range_record : t -> key_type -> key_type -> record_type list
  (** select_range_record d k1 k2 returns a list records present in d whose keys lie 
      between k1 and k2 (inclusive). The list must be sorted in ascending order of the keys.
      You can assume that k1 < k2.**)
      
   val select_by_predicate : t -> (record_type -> bool) -> record_type list
  (** select_by_predicate d f returns a list of all records in d which satisfy the predicate f.
      This list must be sorted in increasing order. If no records satisfy f, the list will be empty **)
  
  val update_record : t -> key_type -> value_type -> unit
  (** update_record d k v updates the record containing key k to have the value v. If the database
      does not contain any such record, it remains unchanged.**)
  
  val update_range_record : t -> key_type -> key_type -> value_type -> unit
  (** update_range_record d k1 k2 v updates all the records in d containing keys 
      between k1 to k2 (inclusive) to have the value v. If the database
      does not contain any such record, it remains unchanged.
      You can assume that k1 < k2.**)
      
  val update_by_predicate: t -> (record_type -> bool) -> value_type -> unit
  (** update_by_predicate d f v updates all records which satisfy the predicate f by changing
      the value to v. Rest of the database remains unchanged. **)
      
  val delete_record : t -> key_type -> unit
  (** delete_record d k deletes the record containing key k in d (if such a record is present),
      leaving the rest of the database unchanged.
  **)
  
  val delete_by_predicate: t -> (record_type -> bool) -> unit
  (** delete_by_predicate d f deletes all records of the database which satisfy the predicate f, 
      leaving the rest of the database unchanged.
  **)
  
  val select_min_record : t -> record_type option
  (** select_min_record d returns the record with the smallest key in the database,
      or None if the database is empty
  **)
  
  val select_max_record : t -> record_type option
  (** select_min_record d returns the record with the largest key in the database,
      or None if the database is empty
  **)
  
  val select_all : t -> record_type list
  (** select_all d returns a list all the records in the database, sorted in increasing
      order according to the keys
  **)

end

```

<!-- grade_id: cell-582552e6e2febff4, points: 0, grade: false, solution: true -->
```OCaml
(* Implement the MakeDatabase functor below. DO NOT CHANGE THE SHARING CONSTRAINTS. *)

exception PrimaryKeyViolation

module MakeDatabase (K: Key) (V: Value) : Database with type key_type = K.t
                                                   with type value_type = V.t
                                                   with type record_type = MakeRecord (K) (V).t
                                                   = struct
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
end

```

```OCaml

module StudentsKey:Key with type t = string = struct
  type t = string
  let compare = String.compare
  let string_of_t k = k 
end

type student_value_type = {name: string; hostel: string; cgpa: float}

module StudentsValue:Value with type t = student_value_type = struct
  type t = student_value_type
  let string_of_t v = v.name ^ ", " ^ v.hostel ^ ", " ^ (string_of_float v.cgpa)
end

module StudentsRecord = MakeRecord (StudentsKey) (StudentsValue)

module StudentsTable = MakeDatabase (StudentsKey) (StudentsValue)

module FacultyKey:Key with type t = string = struct
  type t = string
  let compare = String.compare
  let string_of_t k = k 
end

type faculty_value_type = {name:string; designation:string; salary:float}

module FacultyValue:Value with type t = faculty_value_type = struct
  type t = faculty_value_type
  let string_of_t v = v.name ^ ", " ^ v.designation ^ ", " ^ (string_of_float v.salary)
end

module FacultyTable = MakeDatabase (FacultyKey) (FacultyValue)

module FacultyRecord = MakeRecord (FacultyKey) (FacultyValue)


module CourseKey:Key with type t = string = struct
  type t = string
  let compare = String.compare
  let string_of_t k = k 
end

type course_value_type = {name:string; semester:string; facultyname:string}

module CourseValue:Value with type t = course_value_type = struct
  type t = course_value_type
  let string_of_t v = v.name ^ ", " ^ v.semester ^ ", " ^ v.facultyname
end

module CourseRecord = MakeRecord (CourseKey) (CourseValue)

module CourseTable = MakeDatabase (CourseKey) (CourseValue) 

```

<!-- grade_id: cell-cc53801e0a30ad70, points: 10, grade: true, solution: false -->
```OCaml
(* 10 points *)
let sk1 = "AE19S078";;
let sv1 = {name = "Pearce"; hostel = "Jamuna"; cgpa = 7.57};;
let sd = StudentsTable.create_database ();;
let _ = StudentsTable.insert_record sd sk1 sv1;;
let f (r: StudentsRecord.t) = StudentsRecord.get_key r = "AE19S078";;
let l = StudentsTable.select_by_predicate sd f;;
match l with
| [] -> failwith "impossible"
| x :: xs -> assert (StudentsRecord.string_of_t x = "AE19S078:Pearce, Jamuna, 7.57" && xs = []);;
let f (r: StudentsRecord.t) = let v = StudentsRecord.get_value r in v.hostel = "Mandakini";;
let l = StudentsTable.select_by_predicate sd f;;
match l with
| [] -> ()
| _ -> failwith "impossible" 
let r = StudentsTable.select_record sd sk1;;
match r with
| None -> failwith "Error"
| Some r -> assert (StudentsRecord.string_of_t r = "AE19S078:Pearce, Jamuna, 7.57")

```

<!-- grade_id: cell-e41aa1abb55a7082, points: 30, grade: true, solution: false -->
```OCaml
let ck1 = "CS6508";;
let cv1 = {name="NLP"; semester="EVEN"; facultyname="BVRR"};;
let ck2 = "CS6044";;
let cv2 = {name="Speech Technology"; semester="ODD"; facultyname="RAVI"};;
let ck3 = "CS6422";;
let cv3 = {name="Advanced Wireless Communication and Networks"; semester="EVEN"; facultyname="AYON"};;
let ck4 = "CS6196";;
let cv4 = {name="Stochastic Optimization"; semester = "ODD"; facultyname="KKR"};;
let cd = CourseTable.create_database ();;
let _ = CourseTable.insert_record cd ck1 cv1;;
let _ = CourseTable.insert_record cd ck2 cv2;;
let _ = CourseTable.insert_record cd ck3 cv3;;
let _ = CourseTable.insert_record cd ck4 cv4;;
assert (List.map (fun r -> CourseRecord.string_of_t r) (CourseTable.select_range_record cd "CS6196" "CS6480") 
    = ["CS6196:Stochastic Optimization, ODD, KKR"; "CS6422:Advanced Wireless Communication and Networks, EVEN, AYON"]);;
assert (List.map (fun r -> CourseRecord.string_of_t r) 
        (CourseTable.select_by_predicate cd (fun r -> let v = CourseRecord.get_value r in v.semester="ODD")) = 
        ["CS6044:Speech Technology, ODD, RAVI";"CS6196:Stochastic Optimization, ODD, KKR"]);;
let _ = CourseTable.update_record cd "CS6196" {name="Multi-armed bandits"; semester = "EVEN"; facultyname="BVRR"};;
let _ = CourseTable.update_record cd "CS6521" {name="Speech Technology"; semester="ODD"; facultyname="RAVI"};;
assert (List.map (fun r -> CourseRecord.string_of_t r) (CourseTable.select_all cd) = ["CS6044:Speech Technology, ODD, RAVI";
   "CS6196:Multi-armed bandits, EVEN, BVRR";
   "CS6422:Advanced Wireless Communication and Networks, EVEN, AYON";
   "CS6508:NLP, EVEN, BVRR"] );;
let minr = CourseTable.select_min_record cd in
match minr with
| None -> failwith "impossible"
| Some r -> assert (CourseRecord.string_of_t r = "CS6044:Speech Technology, ODD, RAVI");;
let _ = CourseTable.delete_by_predicate cd (fun r -> let v = CourseRecord.get_value r in v.facultyname="BVRR");;
assert (List.map (fun r -> CourseRecord.string_of_t r) (CourseTable.select_all cd) = ["CS6044:Speech Technology, ODD, RAVI";
   "CS6422:Advanced Wireless Communication and Networks, EVEN, AYON"] );;
let nv = {name="Speech Technology"; semester="EVEN"; facultyname="RAVI"};;
let _ = CourseTable.update_by_predicate cd (fun r -> let k = CourseRecord.get_key r in k = "CS6044") nv;;
assert (List.map (fun r -> CourseRecord.string_of_t r) (CourseTable.select_all cd) = ["CS6044:Speech Technology, EVEN, RAVI";
   "CS6422:Advanced Wireless Communication and Networks, EVEN, AYON"] );;
```