Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you add your code only in the place where it says `YOUR CODE HERE` or "YOUR ANSWER HERE". You must remove the line containing the `raise` command. Don't forget to add your name below:

```OCaml
let name = ""
let rollno = ""
```

## Important notes about grading:

1. All code you submit must compile. Programs that do not compile will receive an automatic zero. If you are having trouble getting your assignment to compile, please contact the TAs or the instructor. 
2. Make sure you do not add any new cells. You must write your code in the designated place only. We use an autograder to grade the submissions, and if your code is not in it's designated place, or you have added a new cell, the autograder may not work correctly. In particular, executing the last cell of the file always adds a new cell at the end, so you must delete this newly added cell.
2. All assignments submitted after the deadline will be considered late, and will consume your grace days. 
3. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and changing it after the deadline will be counted as a late submission.

## Problem 1

Implement predecessor function

```ocaml
pred : int -> int
```

which returns the predecessor of the given integer.

<!-- grade_id: cell-304066634b90fa89, points: 0, grade: false, solution: true -->
```OCaml
let pred x = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-785fb22b90a8c878, points: 10, grade: true, solution: false -->
```OCaml
assert (pred 1 = 0)
```