Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you add your code only in the place where it says `YOUR CODE HERE` or "YOUR ANSWER HERE". You must remove the line containing the `raise` command. Add your name and roll number below. If you have discussed any assignment problem with another student, mention the name/roll number of the student along with the problem number as a comment in the box below (using `(* comment *)`).

```OCaml
let name = ""
let rollno = ""
```

## Important notes about grading:

1. All code you submit must compile. Programs that do not compile will receive an automatic zero. If you are having trouble getting your assignment to compile, please contact the TAs or the instructor. 
2. Make sure you do not add any new cells. You must write your code in the designated place only. We use an autograder to grade the submissions, and if your code is not in it's designated place, or you have added a new cell, the autograder may not work correctly. In particular, executing the last cell of the file always adds a new cell at the end, so you must delete this newly added cell.
3. All assignments submitted after the deadline will be considered late, and will consume your grace days. 
4. **Usage of library functions:** No library function is allowed unless it is explicitly mentioned in the question. Any unauthorized usage of a library function will result in 0 marks for that question.
5. Your code will also be tested on private testcases, with equal weightage for each public and private testcase.

## Problem 1

Implement a function

```ocaml
find_freq : 'a -> 'a list -> int
```

such that `find_freq v l` should return the number of times `v` is present in the list `l`.

<!-- grade_id: cell-42dbd6d0d9c53262, points: 0, grade: false, solution: true -->
```OCaml
let rec find_freq v l = 
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
```

<!-- grade_id: cell-89523cdb60f8ccef, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (find_freq 2 [1;2;2;4;5] = 2)
;assert (find_freq 5 [1;2;3;4] = 0)
```

## Problem 2

Consider a robot which can only move on a straight line, either left or right from its current position. Suppose you are given a set of instructions as a list `l` of strings, where each instruction is either the string `"L"` or `"R"`. If the instruction is `"L"`, then the robot should move one step to the left, and if the instruction is `"R"`, then the robot should move one step to the right. Implement a function count_moves:

```ocaml
count_moves : int -> string list -> int
```

which takes as input the original position of the robot and a list of instructions, and counts the number of moves made by the robot while executing the instructions.

<!-- grade_id: cell-dde487de40e8e6ea, points: 0, grade: false, solution: true -->
```OCaml
let rec count_moves s l = 
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
```

<!-- grade_id: cell-53a730948e0a4b64, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (count_moves 1 ["L";"R"] = 2)
```

## Problem 3

Following on from problem 2, implement a function 

```ocaml
final_position : int -> string list -> int
```

which takes as input the original position of the robot and a list of instructions, and returns the final position of the robot after executing all the instructions. (Hint: `find_freq` from Problem 1 may be handy)

<!-- grade_id: cell-ab8740ed17faebe4, points: 0, grade: false, solution: true -->
```OCaml
let rec final_position s l = 
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
```

<!-- grade_id: cell-0f2c502a584566af, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (final_position 10 ["L";"R"] = 10)
```

## Problem 4

Implement a function 

```ocaml
insert_at : 'a -> int -> 'a list -> 'a list
```

such that `insert_at v i l` should insert `v` at the index `i` in list `l` and return the resultant list. Note that indexing starts from 0. If index `i` exceeds the length of `l`, then `v` should be inserted at the end of the list. The behavior is undefined if `i` $<0$ (i.e. we will not test for such values). 

<!-- grade_id: cell-d741dc96c880b9b5, points: 0, grade: false, solution: true -->
```OCaml
let rec insert_at v i lst =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-198064874b755efe, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (insert_at 1 0 [2] = [1;2]);
assert (insert_at 10 4 [1; 2; 3; 4] = [1;2;3;4;10])
```

## Problem 5

Implement a function

```ocaml
remove : 'a -> 'a list -> 'a list
```

such that `remove v l` should remove every occurrence of `v` from the list `l`, and return the resultant list. Further, you need to ensure that the implementation is tail-recursive. You can use the library function List.rev (which reverses a list) in your implementation. More information on the List module is here: https://ocaml.org/manual/5.2/api/List.html. Note that this is the only library function allowed.

There will be a 50% penalty if your implementation is not tail-recursive.

<!-- grade_id: cell-9edc7695dfbf3748, points: 0, grade: false, solution: true -->
```OCaml
let remove v l =
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")
```

<!-- grade_id: cell-1759d3ed2b1c8ebb, points: 15, grade: true, solution: false -->
```OCaml
(* 15 points *)
assert (remove "d" ["a";"b";"c";"d"] = ["a";"b";"c"]);
assert (remove 2 [1;2;3;2;4;5] = [1;3;4;5])

```

## Problem 6

In this problem, we will implement a method to compute similarity between two lists. This could be used, for example, in plagiarism detection to detect similarity between two documents (by representing a document as a list of strings).

First, we associate each list with a mathematical vector, represented as a list of numbers. The indices of the vector correspond to elements of the list, and the value at each index is how many times that element occurs in the list.

For example, for the lists `l1 = ["a";"b";"c";"d"]` and `l2 = ["d";"d";"d";"b"]`, we first compute vectors of length 4, accounting for all the unique elements in the two lists. The vectors for l1 and l2 would be as follows:

| | `"a"` | `"b"` | `"c"` | `"d"` |
|--- | --- | --- | --- | --- |
| l1 | 1 | 1 | 1 |1|
| l2 | 0 | 1 | 0| 3|

These can be represented as the integer lists `d1 = [1;1;1;1]` and `d2 = [0;1;0;3]`. 

We define the **overlap** between two lists as the dot product of the two list vectors. The dot product is as per the standard definition, i.e. multiply the values of the elements at the same indices in the two vectors and sum them all up. We denote this by d1 $\cdot$ d2. In the above case, the overlap would be 4.

The **similarity score** between two lists is then obtained using the following formula:

$\frac{d1 \cdot d2}{max(||d1||, ||d2||)}$

Here, $||d||$ denotes the magnitude of the vector, obtained by taking the dot product of the vector with itself (i.e. $||d|| = d \cdot d$). Note that the similarity score of a list with itself will be 1, and if two lists have no elements in common, then the similarity score will be 0. In general, the similarity score will be always between 0 and 1. In our running example, the similarity score would be $\frac{4}{max(4,10)} = 0.4$

First, implement a function 

```ocaml
overlap : 'a list -> 'a list -> int
```

which computes the overlap between two lists. Then, use this function to implement

```ocaml
sim_score : 'a list -> 'a list -> float
```

to compute the similarity score. You can use the function `float_of_int : int -> float` to cast an int value to a float value.  Apart from this, you are not allowed to use any other library functions in this problem. Any solution using other library functions will get 0 marks. The functions `find_freq` and `remove` from earlier problems will come in handy.

You can assume that `sim_score` will only be called with non-empty lists.


<!-- grade_id: cell-a49121b6ebd0cbe8, points: 0, grade: false, solution: true -->
```OCaml
let rec overlap l1 l2 = 
  (* YOUR CODE HERE *)
  raise (Failure "Not implemented")

let sim_score l1 l2 = 
(* YOUR CODE HERE *)
raise (Failure "Not implemented")
```

<!-- grade_id: cell-0ec3a97265e6ea89, points: 25, grade: true, solution: false -->
```OCaml
(* 25 points *)
assert (sim_score ["a";"b";"c";"d"] ["d";"d";"d";"b"] = 0.4);
assert (sim_score ["a";"b";"c";"d"] ["b";"a"] = 0.5);
assert (sim_score ["a";"b";"c";"d"] ["b";"a";"d"] = 0.75)


```